# Bonus Assignment 2: Buchi Automata Operations

In this bonus assignment, you will extend your implementation of `FiniteAutomaton` with several useful methods. Each method corresponds to an operation on Buchi automata. The provided methods are helpful for performing analysis and transformations in automata theory.

Be sure to test your implementations thoroughly. Example test cases are provided for each part.

---

```python
import sys
import otter

# try:
#     import otter
# except ImportError:
#     %pip install otter-grader
#     import otter

grader = otter.Notebook("HWB2.ipynb")
```

## Task 1: Reachable States

### Method
```python
def reach(self) -> Set[State]:
        """
        Constructs the sub-automaton of reachable components from initial to accepting states.

        :return: A new finite automaton restricted to reachable components.
        """
```

### Description

This method should return the set of all states that can be reached from any initial state using zero or more transitions.


### Example

If the automaton has transitions:

```
s0 --a--> s1
s1 --b--> s2
```

and initial state is `s0`, and accepting state is `s1`, then `reach()` should return `{s0, s1}`.

---

## Task 2: Closure Under Transitions

### Method
```python
def closure(self, S: Set[State]) -> Set[State]:
    """
    Returns the closure automaton.

    :return: A new NBA A' where L_omega(A') = closure(L_omega(A)).
    """
```

---

## Task 3: Complement Automaton


### Method
```python
def complement(self) -> "FiniteAutomaton":
    """
    Constructs an automaton that accepts the complement of the current language.

    :return: The complement automaton.
    """
```

### Complementing Büchi Automata

Complementing a Büchi automaton is a known difficult problem in automata theory.tly.

To simplify the task for this assignment, we make the following **assumption**:

> The given Büchi automaton represents a **safety property**.

This assumption enables a more straightforward construction:

You are tasked to introduce a new state named `__qfinal__`, which serves as a "sink" or "trap" state. For any state in the original automaton that lacks a transition for a given symbol, a new transition can be added to `__qfinal__` under that symbol. The `__qfinal__` state should then include self-loops for all symbols in the alphabet and be marked as accepting, ensuring that any trace violating the original safety property is accepted by the complement.

---


## Task 4: Union


### Method
```python
def union(self, A: "FiniteAutomaton") -> "FiniteAutomaton":
    """
    Constructs the union of this automaton with another.

    :param A: Another finite automaton.
    :return: A new automaton accepting the union language.
    """
```

### ☯️ Union of Two Büchi Automata

The **union** of two Büchi Automata \( A_1 \) and \( A_2 \) accepts all infinite words accepted by **either** \( A_1 \) or \( A_2 \).

When implementing the union, **you must ensure that the states of the two automata remain distinct**, even if they happen to share the same name (e.g., both contain a state `"q0"`).

To solve this, represent each state in the resulting automaton as a tuple:
```
(state, index)
```
- `state`: the original state name from either automaton
- `index`:
  - `1` if the state comes from `self` (the current automaton),
  - `2` if it comes from `A` (the other automaton passed to the method)

This ensures that `("q0", 1)` and `("q0", 2)` are treated as distinct states in the union, avoiding accidental collisions.

---

## Task 5: Decomposition into Safety and Liveness


### Method
```python
def decompose(self) -> Tuple["FiniteAutomaton", "FiniteAutomaton"]:
    """
    Decomposes the automaton into safe and live components.

    :return: A tuple (A_safe, A_live).
    """
```


### Background

According to the **Decomposition Theorem** in Linear Time Logic, every linear time property can be expressed as the intersection of a **safety property** and a **liveness property**:

\[
P = P_{safety} \cap P_{liveness}
\]

- A **safety property** asserts that "nothing bad happens" — violations can be detected by a finite prefix.
- A **liveness property** asserts that "something good eventually happens" — no finite prefix can conclusively show a violation.

Decomposing an automaton into these two components helps isolate bugs and prove different aspects of correctness.


### Example

```python
P = FiniteAutomaton(
  states={'p', 'y', 'x', 'q'},
  alphabet={'b', 'a'},
  transitions={('x', 'true', 'y'), ('q', 'a', 'q'), ('y', 'true', 'y'), ('p', 'true', 'p'), ('q', 'b', 'p')},
  initial_states={'x', 'q'},
  accepting_states={'p'}
)
```

```python
P_Safe = FiniteAutomaton(
  states={'p', 'q'},
  alphabet={'b', 'a'},
  transitions={('q', 'a', 'q'), ('p', 'true', 'p'), ('q', 'b', 'p')},
  initial_states={'q'},
  accepting_states={'p', 'q'},
)
```

```python
P_Live = FiniteAutomaton(
  states={('p', 2), ('___qfinal___', 2), ('p', 1), ('q', 1), ('q', 2)},
  alphabet={'b', 'a'},
  transitions={(('___qfinal___', 2), 'True', ('___qfinal___', 2)), (('q', 1), 'b', ('p', 1)), (('p', 1), 'true', ('p', 1)), (('q', 2), 'not((b) or (a))', ('___qfinal___', 2)), (('p', 2), 'not((true))', ('___qfinal___', 2)), (('p', 2), 'true', ('p', 2)), (('q', 2), 'a', ('q', 2)), (('q', 1), 'a', ('q', 1)), (('q', 2), 'b', ('p', 2))},
  initial_states={('q', 1), ('q', 2)},
  accepting_states={('___qfinal___', 2), ('p', 1)}
)
```


## Code:

```python
import os
from typing import Set, Tuple, Optional, Union

import networkx as nx
from graphviz import Digraph

State = Union[str, Tuple]
Symbol = str
Transition = Tuple[State, Symbol, State]


class FiniteAutomaton:
    """
    A finite automaton (NFA-style) representation.

    Attributes:
        Q (Set[State]): The set of all states.
        Sigma (Set[Symbol]): The input alphabet (symbols).
        Transitions (Set[Transition]): Transitions labeled with symbols.
        Q0 (Set[State]): The set of initial states.
        F (Set[State]): The set of accepting (final) states.
    """

    def __init__(
        self,
        states: Optional[Set[State]] = None,
        alphabet: Optional[Set[Symbol]] = None,
        transitions: Optional[Set[Transition]] = None,
        initial_states: Optional[Set[State]] = None,
        accepting_states: Optional[Set[State]] = None,
    ) -> None:
        """
        Initializes the finite automaton.

        :param states: A set of states. Defaults to an empty set.
        :param alphabet: A set of input symbols. Defaults to an empty set.
        :param transitions: A set of transitions, each as (state_from, symbol, state_to).
        :param initial_states: A set of initial states. Defaults to an empty set.
        :param accepting_states: A set of accepting states. Defaults to an empty set.
        """
        self.Q: Set[State] = set(states) if states is not None else set()
        self.Sigma: Set[Symbol] = set(alphabet) if alphabet is not None else set()
        self.Transitions: Set[Transition] = set(transitions) if transitions is not None else set()
        self.Q0: Set[State] = set(initial_states) if initial_states is not None else set()
        self.F: Set[State] = set(accepting_states) if accepting_states is not None else set()

    def add_state(self, *states: State) -> "FiniteAutomaton":
        """
        Adds one or more states to the automaton.

        :param states: One or more states to add.
        :return: The updated automaton.
        """
        self.Q.update(states)
        return self

    def add_symbol(self, *symbols: Symbol) -> "FiniteAutomaton":
        """
        Adds one or more symbols to the input alphabet.

        :param symbols: One or more symbols to add.
        :return: The updated automaton.
        """
        self.Sigma.update(symbols)
        return self

    def add_transition(self, *transitions: Transition) -> "FiniteAutomaton":
        """
        Adds one or more transitions to the automaton.

        :param transitions: Transitions in the form (state_from, symbol, state_to).
        :return: The updated automaton.
        :raises ValueError: If states or symbols are not defined in the automaton.
        """
        for transition in transitions:
            if not isinstance(transition, tuple) or len(transition) != 3:
                raise ValueError(f"Invalid transition format: {transition}. Expected (state_from, symbol, state_to).")
            state_from, symbol, state_to = transition
            if state_from not in self.Q or state_to not in self.Q:
                raise ValueError("Transition states must be in the state set.")
            if symbol not in self.Sigma:
                raise ValueError("Transition symbol must be in the alphabet.")
            self.Transitions.add(transition)
        return self

    def add_initial_state(self, *states: State) -> "FiniteAutomaton":
        """
        Adds one or more initial states.

        :param states: States to mark as initial.
        :return: The updated automaton.
        :raises ValueError: If any state is not in the automaton's state set.
        """
        for state in states:
            if state not in self.Q:
                raise ValueError("Initial state must be in the state set.")
            self.Q0.add(state)
        return self

    def add_accepting_state(self, *states: State) -> "FiniteAutomaton":
        """
        Adds one or more accepting (final) states.

        :param states: States to mark as accepting.
        :return: The updated automaton.
        :raises ValueError: If any state is not in the automaton's state set.
        """
        for state in states:
            if state not in self.Q:
                raise ValueError("Accepting state must be in the state set.")
            self.F.add(state)
        return self

    def actions(self, q: State) -> Set[Symbol]:
        """
        Returns the set of symbols that can be used to transition from state q.

        :param q: The state from which to get the actions.
        :return: A set of symbols that can be used to transition from state q.
        """
        return {symbol for (state_from, symbol, state_to) in self.Transitions if state_from == q}

    def reach(self) -> "FiniteAutomaton":
        """
        Constructs the sub-automaton of reachable components from initial to accepting states.

        :return: A new finite automaton restricted to reachable components.
        """
        ...

    def closure(self) -> "FiniteAutomaton":
        """
        Returns the closure automaton.

        :return: A new NBA A' where L_omega(A') = closure(L_omega(A)).
        """
        ...

    def complement(self) -> "FiniteAutomaton":
        """
        Constructs an automaton that accepts the complement of the current language. The current language is assumed to be a safety property.

        :return: The complement automaton.
        """
        ...

    def union(self, A: "FiniteAutomaton") -> "FiniteAutomaton":
        """
        Constructs the union of this automaton with another.

        :param A: Another finite automaton.
        :return: A new automaton accepting the union language.
        """
        ...

    def decompose(self) -> Tuple["FiniteAutomaton", "FiniteAutomaton"]:
        """
        Decomposes the automaton into safe and live components.

        :return: A tuple (A_safe, A_live).
        """
        ...

    def __repr__(self) -> str:
        return (
            f"FiniteAutomaton(\n"
            f"  states={self.Q},\n"
            f"  alphabet={self.Sigma},\n"
            f"  transitions={self.Transitions},\n"
            f"  initial_states={self.Q0},\n"
            f"  accepting_states={self.F}\n"
            f")"
        )

    def to_graphviz(self, name: str = "FA", format: str = "pdf", view: bool = True, filename: Optional[str] = None) -> None:
        """
        Visualizes the finite automaton using Graphviz.

        :param name: Name of the Graphviz graph.
        :param format: Output file format (e.g., 'pdf', 'png', 'svg').
        :param view: Whether to open the output file after rendering.
        :param filename: Optional filename (without extension). If None, uses name.
        """
        dot = Digraph(name=name, format=format)
        dot.attr(rankdir="LR", size="8,5")

        # Invisible initial node for pointing to real initial states
        dot.node("", shape="none")

        for state in self.Q:
            shape = "doublecircle" if state in self.F else "circle"
            dot.node(str(state), shape=shape)

        for q0 in self.Q0:
            dot.edge("", str(q0))

        for (s, a, t) in self.Transitions:
            dot.edge(str(s), str(t), label=str(a))

        output_file = filename if filename is not None else name
        dot.render(output_file, view=view, cleanup=True)
```

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> def test_reach1():
...     p = FiniteAutomaton(states={'s0', 's1', 's2'}, alphabet={'a'}, transitions={('s0', 'a', 's1'), ('s1', 'a', 's2')}, initial_states={'s0'}, accepting_states={'s1'})
...     assert 's2' not in p.reach().Q
>>> def test_reach2():
...     p = FiniteAutomaton(states={'q', 'p', 'x', 'y'}, alphabet={'a', 'b'}, transitions={('q', 'a', 'q'), ('q', 'b', 'p'), ('p', 'true', 'p'), ('x', 'true', 'y'), ('y', 'true', 'y')}, initial_states={'q', 'x'}, accepting_states={'p'})
...     p = p.reach()
...     assert 'x' not in p.Q
...     assert 'y' not in p.Q
...     assert 'x' not in p.Q0
...     assert p.F == {'p'}
>>> def test_closure():
...     p = FiniteAutomaton(states={'q', 'p', 'x', 'y'}, alphabet={'a', 'b'}, transitions={('q', 'a', 'q'), ('q', 'b', 'p'), ('p', 'true', 'p'), ('x', 'true', 'y'), ('y', 'true', 'y')}, initial_states={'q', 'x'}, accepting_states={'p'})
...     p = p.reach()
...     clp = p.closure()
...     assert clp.Q == p.Q
...     assert clp.F == p.Q
>>> def test_complement():
...     p = FiniteAutomaton(states={'q', 'p', 'x', 'y'}, alphabet={'a', 'b'}, transitions={('q', 'a', 'q'), ('q', 'b', 'p'), ('p', 'true', 'p'), ('x', 'true', 'y'), ('y', 'true', 'y')}, initial_states={'q', 'x'}, accepting_states={'p'})
...     p = p.reach()
...     clp = p.closure()
...     cclp = clp.complement()
...     assert cclp.F == {'___qfinal___'}
...     assert cclp.Q0 == p.Q0
...     assert cclp.Q == p.Q | {'___qfinal___'}
>>> def test_union():
...     p1 = FiniteAutomaton(states={'q1', 'p1'}, alphabet={'a', 'b'}, transitions={('q1', 'a', 'p1'), ('p1', 'b', 'q1')}, initial_states={'q1'}, accepting_states={'p1'})
...     p2 = FiniteAutomaton(states={'q2', 'p2'}, alphabet={'a', 'b'}, transitions={('q2', 'a', 'p2'), ('p2', 'b', 'q2')}, initial_states={'q2'}, accepting_states={'p2'})
...     union = p1.union(p2)
...     assert len(union.Q) == 4
...     assert len(union.Transitions) == 4
...     assert union.Q0 == {('q1', 1), ('q2', 2)}
...     assert union.F == {('p1', 1), ('p2', 2)}
>>> def test_decompose():
...     p = FiniteAutomaton(states={'q', 'p', 'x', 'y'}, alphabet={'a', 'b'}, transitions={('q', 'a', 'q'), ('q', 'b', 'p'), ('p', 'true', 'p'), ('x', 'true', 'y'), ('y', 'true', 'y')}, initial_states={'q', 'x'}, accepting_states={'p'})
...     p = p.reach()
...     (A_safe, A_live) = p.decompose()
...     assert A_safe.Q == {'p', 'q'}
...     assert A_live.Q == {('p', 2), ('___qfinal___', 2), ('p', 1), ('q', 1), ('q', 2)}
...     assert A_live.F == {('___qfinal___', 2), ('p', 1)}
>>> test_reach1()
>>> test_reach2()
>>> test_closure()
>>> test_complement()
>>> test_union()
>>> test_decompose()
```
